// intro2.rs
//
// Make the code print a greeting to the world.
//
// Execute `rustlings hint intro2` or use the `hint` watch subcommand for a
// hint.


// fn main() {
//     println!("Hello {}!","world");
// }

// I AM NOT DONE

// fn main() {
//     println!("Hello {var}!",var="world");
// }


fn main() {
    println!("Hello {0}\nHELLO {1}!","world","WORLD");
}